# Submit your work!

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and assign your slack id to a `slack_id` variable in the cell bellow.
Example:

```python
slack_id = "UTS63FC02"
```

Then, in the last cell, do the following steps in this order:
1. Make sure the line `#grade_submit(...)` is a comment (starts with the `#` character)
1. Save the notebook
1. Uncomment the line `#grade_submit(...)` by removing the `#` character
1. Run the last cell of the notebook
1. You should get an output like `Success {"id":102,"created":"2021-12-24T18:38:50.198035Z","slackid":"UTS63FC02","learning_unit":0,"exercise_notebook":1,"score":20.0}`


<!-- grade_id: cell-04995e54d9b0e3c4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
slackid = "UTS63FC02"
### END SOLUTION
# slackid =
```

<!-- grade_id: cell-1d7c7c44c66fab01, points: 20, grade: true, solution: false -->
```python
from submit import grade_submit

assert slackid is not None
#grade_submit(notebook_name='Exercise notebook.ipynb', learning_unit=0, exercise_notebook=1, slackid=slackid)
```

Output:
```
Success {"id":102,"created":"2021-12-24T18:38:50.198035Z","slackid":"UTS63FC02","learning_unit":0,"exercise_notebook":1,"score":20.0}
```